// range2
// Make me compile!
//
// Iterate a map's keys and values with range.
//
// I AM NOT DONE
package main

import "fmt"

func main() {
	phoneBook := map[string]string{
		"Ana":  "+01 101 102",
		"John": "+01 333 666",
	}

	for := range {
		fmt.Printf("%s has the %s phone\n", name, phone)
	}
}
